<a href="https://colab.research.google.com/github/laxmipranavimallekedi/HDS5210_Instructor/blob/main/week04/week04-assignment_03.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 4 Assignment

See: _McKinney 3.1_ and [Python Documentation](https://docs.python.org/3/tutorial/datastructures.html) section 5 on data structures.

Modules 13 through 15 have all been about different built-in data types that can be used to create larger and more complex data structures.  We learned about tuples, lists, dictionaries, and how to access and process that data in new ways.

**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 15.1 Create a sorted list of names

Take the list of names below (some are first and last, others are first, middle, and last) and produce a final list of full names.

```
sort_names([
  ['Adam', 'Doyle'],
  ['Kermit','the','Frog'],
  ['Donald','Duck']])

['Adam Doyle', 'Donald Duck', 'Kermit the Frog']
```

**Tip:** If you have a list of lists and want to sort it, Python can do that. It will simply sort using the first item in each of the lists inside your list.

```python
names = [['Adam', 'Doyle'],
  ['Kermit','the','Frog'],
  ['Donald','Duck']]
```

```python
def name_sorter(name_list):
    # join names
    full_names = [" ".join(name) for name in name_list ]
    #sorting names
    full_names.sort()
    return full_names
```

```python
names = [['Adam', 'Doyle'],
  ['Kermit','the','Frog'],
  ['Donald','Duck']]

assert name_sorter(names) == ["Adam Doyle", "Donald Duck", "Kermit the Frog"]
```

### 15.2 Convert a dictionary to a list

In this exercise, we're going to take a dictionary that has families and their names, and convert that into a list of names.  The dictionary is structured such that the key is the family name, and the value for that key is a list of first names.

```
{
  'Doyle': ['Adam', 'Oliver', 'Xavier'],
  'Duck': ['Donald', 'Daisy', 'Howard'],
  'Mouse': ['Mickey', 'Minnie', 'Mighty']
}

[['Doyle', 'Adam'],
 ['Doyle', 'Oliver'],
 ['Doyle', 'Xavier'],
 ['Duck', 'Donald'],
 ['Duck', 'Daisy'],
 ...
```

```python
families = {
  'Doyle': ['Adam', 'Oliver', 'Xavier'],
  'Duck': ['Donald', 'Daisy', 'Howard'],
  'Mouse': ['Mickey', 'Minnie', 'Mighty']
}

```

```python
def dict_to_list(families):
    # start by creating an empty list
    output = []
    # loop through the key value pairs in the input dict
    for name,family_list in families.items():
    #loop through each family in family_list
     for member in family_list:
    #add the
      output.append([name,member])

    return output
```

```python
families = {
  'Doyle': ['Adam', 'Oliver', 'Xavier'],
  'Duck': ['Donald', 'Daisy', 'Howard'],
  'Mouse': ['Mickey', 'Minnie', 'Mighty']
}

dict_to_list({'Duck': ['Hewey','Dewey','Louie'],
              'Mouse': ['Mickey','Minney']})
```

Output:
```
[['Duck', 'Hewey'],
 ['Duck', 'Dewey'],
 ['Duck', 'Louie'],
 ['Mouse', 'Mickey'],
 ['Mouse', 'Minney']]
```

### 20.3 Convert a list to a dictionary

Take the list of lists below and convert it into a dictionary where the entries in the dictionary are keyed off of whatever values are in column 0, and the dictionary values are the list of all the values that appear with that key from the input list.

<!-- grade_id: average-answer, points: 0, grade: false, solution: true -->
```python
from os import name
def list_to_dict(names):

    names = [['Doyle', 'Adam'],
 ['Doyle', 'Oliver'],
 ['Doyle', 'Xavier'],
 ['Duck', 'Donald'],
 ['Duck', 'Daisy'],
 ['Mouse', 'Mickey'],
 ['Mouse', 'Minnie'],
 ['Mouse', 'Mighty']
 ]

    # To help you, I've provided a blank dictionary named `families`.
    families = {}

    # Your code should loop through `names` and do something to build
    for member in names:
       family_name = member[0]
       name = member[1]
       if family_name not in families.keys():
          families[family_name]=[]
       families[family_name].append(name)

    # Put your solution here

    return families

```

```python
list_to_dict(names)
```

Output:
```
{'Doyle': ['Adam', 'Oliver', 'Xavier'],
 'Duck': ['Donald', 'Daisy'],
 'Mouse': ['Mickey', 'Minnie', 'Mighty']}
```

<!-- grade_id: average-tests, points: 5, grade: true, solution: false -->
```python
answer = {
 'Doyle':    ['Adam', 'Oliver', 'Xavier'],
 'Duck':    ['Donald', 'Daisy'],
 'Mouse':   ['Mickey', 'Minnie', 'Mighty']
}

assert (list_to_dict(names) == answer)
```

### 20.4 Join using a dictionary

We have a list of patients, diagnosis, and length of stay.  We also have a dictionary that contains diagnosis and average length of stay.  Produce an output list that lists the patient and an indicator if the patient's stay was 'too long', 'too short', 'just right'

<!-- grade_id: cell-f3813e33810dc898, points: 0, grade: false, solution: true -->
```python
def patient_los(avg_los, patients):

    conclusions = []
    for patient in patients:
        name = patient[0]
        diagnosis = patient[1]
        length_of_stay = patient[2]
        if length_of_stay > avg_los[diagnosis]:
            conclusions.append([name,"too long"])
        elif length_of_stay < avg_los[diagnosis]:
            conclusions.append([name,"too short"])
        else:
            conclusions.append([name,"just right"])


    return conclusions
```

```python
avg_los = {
    "Hemolytic jaundice and perinatal jaundice" : 2,
    "Medical examination/evaluation" : 3.2,
    "Liveborn" : 3.2,
    "Trauma to perineum and vulva" : 2.1,
    "Normal pregnancy and/or delivery" : 2,
    "Umbilical cord complication" : 2.1,
    "Forceps delivery" : 2.2,
    "Administrative/social admission" : 4.2,
    "Prolonged pregnancy" : 2.4,
    "Other complications of pregnancy" : 2.5
}

patients = [
    ['Doyle', 'Medical examination/evaluation', 1.1],
    ['Doyle', 'Other complications of pregnancy', 3.3],
    ['Jones', 'Liveborn', 3.2],
    ['Ashbury', 'Forceps delivery', 2.0]
]

patient_los(avg_los, patients)
```

Output:
```
[['Doyle', 'too short'],
 ['Doyle', 'too long'],
 ['Jones', 'just right'],
 ['Ashbury', 'too short']]
```

<!-- grade_id: cell-2af0a520bb56ee21, points: 10, grade: true, solution: false -->
```python
conclusions = [['Doyle', 'too short'],
    ['Doyle', 'too long'],
    ['Jones', 'just right'],
    ['Ashbury', 'too short']]

assert(patient_los(avg_los, patients) == conclusions)
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week04/week04_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

